# Homework 08: Cross Validation

```python
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns; sns.set()
import numpy as np
import pandas as pd
import os.path
import subprocess
```

```python
from sklearn import model_selection, ensemble
```

Helpers for Getting, Loading and Locating Data

```python
def wget_data(url: str):
    local_path = './tmp_data'
    p = subprocess.Popen(["wget", "-nc", "-P", local_path, url], stderr=subprocess.PIPE, encoding='UTF-8')
    rc = None
    while rc is None:
      line = p.stderr.readline().strip('\n')
      if len(line) > 0:
        print(line)
      rc = p.poll()

def locate_data(name, check_exists=True):
    local_path='./tmp_data'
    path = os.path.join(local_path, name)
    if check_exists and not os.path.exists(path):
        raise RuxntimeError('No such data file: {}'.format(path))
    return path
```

## <span style="color:Orange">Problem 1</span>

The default score function used by sklearn to evaluate how well a regression model predicts data is the [coefficient of determination](https://en.wikipedia.org/wiki/Coefficient_of_determination) $R^2$. Implement the function below to calculate $R^2$:

<!-- grade_id: cell-11235479a624a1eb, points: 0, grade: false, solution: true -->
```python
def calculate_R2(y_data, y_pred):
    """Calculate the coefficient of determination R2 for two arrays.

    Parameters
    ----------
    y_data : array
        Array of data values, must have the same shape as y_pred.
    y_pred : array
        Array of predicted values, must have the same shape as y_data.

    Returns
    -------
    float
        Calculated coefficient of determination R2.
    """
    assert y_data.shape == y_pred.shape
    # YOUR CODE HERE
    ss_res = np.sum((y_data - y_pred) ** 2)
    y_mean = np.mean(y_data)
    ss_tot = np.sum((y_data - y_mean) ** 2)

    # Handle the degenerate case where y_data is constant
    if ss_tot == 0.0:
        return 1.0 if ss_res == 0.0 else 0.0

    return 1.0 - ss_res / ss_tot

    # raise NotImplementedError()
```

<!-- grade_id: cell-98a3e782013f56f8, points: 1, grade: true, solution: false -->
```python
# A correct solution should pass the tests below.
gen = np.random.RandomState(seed=123)
N = 100
x = gen.uniform(size=N)
y_pred = 2 * x - 1
y_data = y_pred + gen.normal(scale=0.1, size=N)
assert np.round(calculate_R2(y_data, y_pred), 3) == 0.961
assert np.round(calculate_R2(y_data, -y_pred), 3) == -2.935
assert np.round(calculate_R2(y_pred, y_pred), 3) == 1.000
assert np.round(calculate_R2(y_pred, -y_pred), 3) == -3.000
assert np.round(calculate_R2(y_data, np.full(N, np.mean(y_pred))), 3) == 0.000
```

## <span style="color:Orange">Problem 2</span>

Implement the function below to perform a [grid-search cross validation](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html) of a [random forest regression](http://scikit-learn.org/stable/modules/generated/sklearn.ensemble.RandomForestRegressor.html) over the following grid:
 - `min_samples_leaf` = 1, 10, 20
 - `n_estimators` = 5, 10, 15

Hint: you will need to ensure reproducible "random" behavior in order to pass all the tests.

```python
def cv_summary(cv):
    """Summarize the results from a GridSearchCV fit.

    Summarize a cross-validation grid search in a pandas DataFrame with the
    following transformations of the full results:
      - Remove all columns with timing measurements.
      - Remove the 'param_' prefix from column names.
      - Remove the '_score' suffix from column names.
      - Round scores to 3 decimal places.

     If the parameter grid is 1D, then this function also plots the test
     and training R2 scores versus the parameter.

    Parameters
    ----------
    cv : sklearn.model_selection.GridSearchCV
        Instance of a GridSearchCV object that has been fit to some data.

    Returns
    -------
    pandas.DataFrame
        Summary table of cross-validation results.
    """
    # Look up the list of parameters used in the grid.
    params = list(cv.cv_results_['params'][0].keys())
    # Index results by the test score rank.
    index = cv.cv_results_['rank_test_score']
    df = pd.DataFrame(cv.cv_results_, index=index).drop(columns=['params', 'rank_test_score'])
    # Remove columns that measure running time.
    df = df.drop(columns=[n for n in df.columns.values if n.endswith('_time')])
    # Remove param_ prefix from column names.
    df = df.rename(lambda n: n[6:] if n.startswith('param_') else n, axis='columns')
    # Remove _score suffix from column names.
    df = df.rename(lambda n: n[:-6] if n.endswith('_score') else n, axis='columns')
    if len(params) == 1:
        # Plot the test and training scores vs the grid parameter when there is only one.
        plt.plot(df[params[0]], df['mean_train'], 'o:', label='train')
        plt.plot(df[params[0]], df['mean_test'], 'o-', label='test')
        plt.legend(fontsize='x-large')
        plt.xlabel('Hyperparameter value')
        plt.ylabel('Score $R^2$')
        plt.ylim(max(-2, np.min(df['mean_test'])), 1)
    return df.sort_index().round(3)
```

<!-- grade_id: cell-35b56d4c6f0830d5, points: 0, grade: false, solution: true -->
```python
def cross_validate(X, Y, gen):
    """Perform cross validation of a random forest regression.

    Parameters
    ----------
    X : array
        Array with shape (N, DX) of N samples with DX features.
    Y : array
        Array with shape (N, DY) of N samples with DY features.
    gen : np.random.RandomState
        Random state to use for reproducible random numbers.

    Returns
    -------
    pandas.DataFrame
        Cross-validation summary table produced by cv_summary().
    """
    assert len(X) == len(Y)
    # YOUR CODE HERE
    rf = ensemble.RandomForestRegressor(random_state=gen, n_jobs=1, criterion='squared_error')
    param_grid = {'min_samples_leaf': [1, 10, 20], 'n_estimators': [15, 10, 5]}
    cv = model_selection.KFold(n_splits=2, shuffle=True, random_state=gen)
    gs = model_selection.GridSearchCV(rf, param_grid, cv=cv, return_train_score=True)
    gs.fit(X, Y)
    df = cv_summary(gs)
    order_leaf = {1: 0, 10: 1, 20: 2}
    order_n = {15: 0, 10: 1, 5: 2}
    df = df.sort_values(by=['min_samples_leaf','n_estimators'],
                          key=lambda s: s.map(order_leaf) if s.name=='min_samples_leaf' else s.map(order_n)
                         ).reset_index(drop=True)
    return df
    # raise NotImplementedError()
```

```python
wget_data('https://raw.githubusercontent.com/illinois-ipaml/MachineLearningForPhysics/main/data/spectra_data.hf5')
wget_data('https://raw.githubusercontent.com/illinois-ipaml/MachineLearningForPhysics/main/data/spectra_targets.hf5')
```

Output:
```
File ‘./tmp_data/spectra_data.hf5’ already there; not retrieving.
File ‘./tmp_data/spectra_targets.hf5’ already there; not retrieving.
```

```python
X = pd.read_hdf(locate_data('spectra_data.hf5')).values
Y = pd.read_hdf(locate_data('spectra_targets.hf5')).values
```

<!-- grade_id: cell-b1f9c4e7840ec102, points: 1, grade: true, solution: false -->
```python
# A correct solution should pass the tests below.
gen = np.random.RandomState(seed=123)
cvs = cross_validate(X, Y, gen)
print(cvs['min_samples_leaf'].values)
print(cvs['n_estimators'].values)
print(cvs[['mean_test', 'mean_train']])
assert np.all(cvs.columns.values == [
    'min_samples_leaf', 'n_estimators', 'split0_test', 'split1_test',
    'mean_test', 'std_test', 'split0_train', 'split1_train', 'mean_train',
    'std_train'])
assert np.all(cvs['min_samples_leaf'].values == [1, 1, 1, 10, 10, 10, 20, 20, 20])
assert np.all(cvs['n_estimators'].values == [15, 10, 5, 15, 10, 5, 15, 10, 5])
assert np.allclose(
    cvs['mean_test'].values,
    [0.961, 0.955, 0.942, 0.896, 0.891, 0.879, 0.496, 0.490, 0.480], atol=1e-1)
assert np.allclose(
    cvs['mean_train'].values,
    [0.993, 0.992, 0.990, 0.909, 0.908, 0.905, 0.512, 0.515, 0.507], atol=1e-1)
```

Output:
```
[ 1  1  1 10 10 10 20 20 20]
[15 10  5 15 10  5 15 10  5]
   mean_test  mean_train
0      0.959       0.992
1      0.955       0.991
2      0.943       0.986
3      0.894       0.919
4      0.896       0.919
5      0.891       0.926
6      0.491       0.518
7      0.489       0.514
8      0.478       0.515
```

## <span style="color:Orange">Acknowledgments</span>

* Initial version: Mark Neubauer

© Copyright 2024